```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import pandas as pd
import numpy as np
import numpy.testing as nt
```

# Optimisation II - Lab 1


## Instructions

* Carefully read all the instructions.
* Do **NOT** change the file name/format upon submission. Please do not add or remove cells from the notebook. 


* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!


----
This lab material has been created by Dr Matthew Woolway and has been edited by Krupa Prag. 

## Section I: Warm-up Exercises

Complete the following warm-up tasks. 

### Question 1 [1 Mark]

Write a function that will convert a decimal number to a binary number. The function must be able to handle a decimal of any length. Do NOT use any built-in functions to complete this task. 

<!-- grade_id: cell-9df3f56b7fca0a7a, points: 0, grade: false, solution: true -->
```python
def dec_to_bin(n):
    """
    Inputs
        n: a scalar real value
    
    Output
        s: a binary string equal to n
    """
    # YOUR CODE HERE
    # binary_num = []
    
    # while n > 0:
    #     rem = n % 2
    #     binary_num.append(rem)
    #     n = n // 2
    # binary_num.reverse()
    # output = "".join([str(i) for i in binary_num])
    
    list = []
    while (n > 0):
        list.append(n%2)
        n = n//2
    list.reverse()
    listToStr = ''.join([str(elem) for elem in list])
    return(listToStr)
    
    # return output
        
    raise NotImplementedError()
```

<!-- grade_id: cell-41f040e25a803077, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
nt.assert_string_equal('1101', dec_to_bin(13))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 2 [1 Mark]

A string is simply an ordered collection of symbols selected from some alphabet and formed into a word; the length of a string is the number of symbols that it contains. An example of a length 21 DNA string (whose alphabet contains the symbols 'A', 'C', 'G', and 'T') is "ATGCTTCAGAAAGGTCTTACG." Write a function that takes as input, a vector  $s$  and returns a vector of four integers counting the respective number of times that the symbols 'A', 'C', 'G', and 'T' occur in  $s$.

<!-- grade_id: cell-6a521f3c5be8821b, points: 0, grade: false, solution: true -->
```python
def dna(s):
    """
    Inputs
    s: a string of DNA sequence
    
    Outputs
    occ_seq: a list of the number of times A, C, G and T occur (in this order)
    """
    # YOUR CODE HERE
    counts = [0,0,0,0]
    
    for i in range(len(s)):
        if s[i] == 'A':
            counts[0] += 1
        if s[i] == 'C':
            counts[1] += 1
        if s[i] == 'G':
            counts[2] += 1
        if s[i] == 'T':
            counts[3] += 1     
    
    return counts
    
    raise NotImplementedError()
```

<!-- grade_id: cell-acd8c87c71d13c6f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
s = 'ATGCTTCAGAAAGGTCTTACG'
nt.assert_array_equal([6, 4, 5, 6], dna(s), err_msg='incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 3 [1 Mark]

Write a function that takes a matrix  $M$  as input and returns a matrix that contains only those elements of  $M$  that are in odd rows and columns.

<!-- grade_id: cell-9ede6d2acfe528d8, points: 0, grade: false, solution: true -->
```python
def matrix_extraction(M):
    """
    Inputs:
        M: a random numpy matrix
    
    Outputs:
        N: a matrix with the odd rows and columns of M only
    """
    # YOUR CODE HERE
    cols = M.shape[1]
    rows = M.shape[0]
    
    c = []
    r = []
    
    for i in range(rows):
        if i % 2 != 0:
            r.append(i)
    
    for i in range(cols):
        if i % 2 != 0:
            c.append(i)
    
    M = np.delete(M,r,0)
    M = np.delete(M,c,1)
        
    return M
    raise NotImplementedError()
```

<!-- grade_id: cell-058e4cee17fd3825, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
M = np.array([[0.29977941, 0.05186872, 0.57587171, 0.56806655, 0.38424249],
             [0.25568574, 0.53849314, 0.76387274, 0.29669909, 0.68029403],
             [0.9338883,  0.44839289, 0.48052527, 0.60921226, 0.89123473],
             [0.76114271, 0.61475746, 0.04465618, 0.85599524, 0.29450567],
             [0.43419249, 0.04936524, 0.14366823, 0.30798834, 0.91016638]])
N = np.array([[0.29977941, 0.57587171, 0.38424249],
             [0.9338883,  0.48052527, 0.89123473],
             [0.43419249, 0.14366823, 0.91016638]])
nt.assert_array_almost_equal(N, matrix_extraction(M), err_msg='incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 4 [1 Mark]

Write a function that computes how much money we have. It takes a single input argument, a vector  $c$, whose elements specify the number of five cent, 10 cent, 20 cent and 50 cent coins we have (in the order listed here). The output should be a scalar value of the total in Rands - not cents.

<!-- grade_id: cell-51e6c91fccb0f4f2, points: 0, grade: false, solution: true -->
```python
def rand_compute(c):
    """
    Inputs
        c: a list of the number of 5, 10, 20 and 50 cent coins
        
    Outputs
        total: the scalar Rand value of all the coins in the list
    """
    # YOUR CODE HERE
    five = 1/20
    ten = 1/10
    twenty = 1/5
    fifty = 1/2
    r = (c[0]*five) + (c[1]*ten) + (c[2]*twenty) + (c[3]*fifty)
    
    return r
    raise NotImplementedError()
```

<!-- grade_id: cell-94cb627b5e119e90, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
c = [0,1,0,2]
nt.assert_equal(1.1, rand_compute(c), err_msg= 'incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 5 [1 Mark]

Write a function that that takes in two vectors of items and returns a vector with the symmetric difference of the two vectors of items. For example if lists $A = [1,8,9,7,14]$ and $B = [5,9,6,14,1]$ are the inputs, then return the symmetric difference of these lists which is $S=[8,7,5,6]$. 

<!-- grade_id: cell-23227cecdd1a04d0, points: 0, grade: false, solution: true -->
```python
def set_union(A, B):
    """
    Inputs
        A: a list of items
        B: a list of items

        
    Outputs
        S: the symmetric difference of the lists A and B
    """
    # YOUR CODE HERE
    A = set(A)
    B = set(B)
    
    A_B = set(A.difference(B))
    B_A = set(B.difference(A))
    
    sym_diff = A_B.union(B_A)
    
    # difference = list(set(A).symmetric_difference(set(B)))
    
    return sym_diff
    raise NotImplementedError()
```

<!-- grade_id: cell-c093ced4acb9eb4c, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
A = ['cat', 'sat', 'bat', 'mat']
B = ['bat', 'dog', 'pet', 'cat']
nt.assert_equal(set(['pet', 'sat', 'dog', 'mat']), set(set_union(A, B)), err_msg= 'incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

## Section II: Main Exercise

### Question 1

Write a function which randomly samples a local neighbourhood. This neighbourhood maybe circular or square in nature. The function takes as input, a vector of (x, y) coordinate points, i.e. $[x, y]$. A scalar, $r$, a radial value, a scalar, $n$, the number of samples and scalar, $s$, the shape to explore ($s = 1$ - a circular radius and $s =2$ a square neighbourhood). The function should return an array containing the coordinates of the randomly generated samples. 

Finally, plot these coordinates and its "origin" on the same set of axes to see your function in action. Explore with different values for your inputs.

<!-- grade_id: cell-6d9b35ecc977fe99, points: 0, grade: false, solution: true -->
```python
def neighbourhood(c, r, n, s):
    # YOUR CODE HERE
    
    # xpoints = []
    # ypoints = []
    # figure, ax = plt.subplots()
        
    # if s==1 :    
    #     circle = plt.Circle(c, r, fill=False)
    #     ax.add_patch(circle)
        
    #     for i in range(n):
    #         # rad, theta = [sqrt(np.random.randint(0,r))*sqrt(r), 2*pi*np.random.random()]
    #         theta = 2 * pi * np.random.random()
    #         x = c[0] + r * cos(theta) 
    #         nx = c[0] - r * cos(theta)
    #         y = c[1] + r * sin(theta)
    #         ny = c[1] - r * sin(theta)
    #         xpoints.append(np.random.uniform(x, nx))
    #         ypoints.append(np.random.uniform(y, ny))

        
    # else:
    #     l = 2*r
    #     center = (c[0] - r, c[1] - r)
    #     rectangle = plt.Rectangle(center, l, l, fill=False)
    #     ax.add_patch(rectangle)
        
    #     xpoints = list(np.random.uniform(c[0] - r, c[0] + r, n))
    #     ypoints = list(np.random.uniform(c[1] - r, c[1] + r, n))
    
    # points = []
    
    # for i in range(n):
    #     if xpoints[i] <= (c[0] + r) and xpoints[i] >= (c[0] - r):
    #        if ypoints[i] <= (c[1] + r) and ypoints[i] >= (c[1] - r):
    #            points.append((xpoints[i], ypoints[i]))
               
    # ax.set_aspect(1)
    # plt.grid()
    # ax.plot()
    # plt.plot(c[0], c[1], "o", 7, markerfacecolor="red")
    # plt.scatter(xpoints, ypoints, 4)
    # plt.show()
    
    # return points
    
    # x, y = c[0], c[1]
    # xpoints = np.empty(n)
    # ypoints = np.empty(n)
    
    # for i in range(n):
    #     theta = 2 * pi * np.random.random(1)
    #     R = r * np.random.random(1)
    
    #     xpoints[i] = x + R * cos(theta)
    #     ypoints[i] = y + R * sin(theta)
           
    # return [xpoints, ypoints]

    x, y = c[0], c[1]
    xpoints = np.empty(n)
    ypoints = np.empty(n)
    
    for i in range(n):
        theta = 2 * pi * np.random.random(1)
        R = r * np.random.random(1)
        
        xpoints[i] = x + R * cos(theta)
        ypoints[i] = y + R * sin(theta)
        
    return [xpoints, ypoints]
        
    
    raise NotImplementedError()
```

```python
# Question 1 is not graded. You see run the below code to see if your code correctly generates the neigbourhoods
# c = [np.random.randint(10), np.random.randint(10)]
# r = 5
# n = 1000
# s = 1
# neighbourhood(c, r, n, s)

c = [np.random.randint(10), np.random.randint(10)]
r = np.random.randint(10)
n = 20
s = 1
xtest, ytest = neighbourhood(c, r, n, s)
for i in range(n):
    pt = np.array([xtest[i], ytest[i]])
    dist = np.linalg.norm(pt - c, 2)
    if dist > r:
        raise AssertionError('Point exceeds max euclidean distance')
print('Test case passed!')



```

Output:
```
Test case passed!
```

### Question 2 [1 Mark]

Write a function which takes in as input an unordered vector of values. Your function should return a row vector ($1 \times 3$) containing the minimum value, median value and maximum value (in this order). Do not use the builtin sort function. Look up sorting algorithms, *bubble sort* for example.

<!-- grade_id: cell-39600d74da63a6c1, points: 0, grade: false, solution: true -->
```python
def lst_arrange(x):
  """
  Inputs:
      x: an array of random integers
  
  Output:
      vals: an array of 3 elements, in the order of min, median and max values of the input
  """
  # YOUR CODE HERE
  
  low = 0
  high = len(x)-1
  
  print(x)
  
  for i in range(0, len(x) - 1):
    for j in range(0,len(x) - i - 1):
      if x[j] > x[j + 1]:
        (x[j], x[j + 1]) = (x[j+1], x[j]) 
  
  print(x)
  # quick_sort(x,low, high)
  
  final = []
  final.append(x[0])
  
  if len(x) % 2 != 0:
    final.append(x[len(x)//2])
  else:
    median = (x[len(x)//2] + x[len(x)//2 + 1])/2 
    final.append(median)
    
  final.append(x[len(x)-1])
  
  return final
  raise NotImplementedError()


# def partition(array, low, high):
#   pivot = array[high]
  
#   i = low - 1
  
#   for j in range(low, high):
#     if array[j] <= pivot:
#       i = i + 1
#       (array[i], array[j]) = (array[j], array[i])
  
#   (array[i + 1], array[high]) = (array[high], array[i + 1])
  
#   return i + 1
  
  
# def quick_sort(array, low, high):
#   if low < high:
#     pi = partition(array, low, high)

#     quick_sort(array, low, pi - 1)
    
#     quick_sort(array, pi + 1, high)
      
```

<!-- grade_id: cell-1ca35137aef43e0e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
y = np.random.randint(1, 10, 5)
nt.assert_array_almost_equal([np.min(y), np.median(y), np.max(y)], lst_arrange(y), err_msg= 'incorrect function')
print('All Tests Passed!!!')
```

Output:
```
[4 8 2 7 8]
[2 4 7 8 8]
All Tests Passed!!!
```

### Question 3

Write a function which takes as input an array of points, i.e. $[(x_1, y_1), (x_2, y_2), \ldots]$. It should return as output, a smaller array containing the points of the Convex Hull vertices (ensure no duplicates).  Some algorithms which can find the convex hull are *Andrew's monotone chain algorithm* and *gift wrapping*, there are others - choose whichever one you like.

Note: Colinear points should be omitted from the returned vertices. 


*You may plot this Convex Hull and the original set of points on the same axes to determine if your function is doing as instructed.*

<!-- grade_id: cell-babaf6656da06547, points: 0, grade: false, solution: true -->
```python
# def l_index(points):
#     min = 0
#     for i in range(1,len(points)):
#         if points[i][0] <= points[min][0]:
#             min = i
#     return min
 
# def orientation(p, q, r):
#     val = (q[1] - p[1]) * (r[0] - q[0]) - (q[0] - p[0]) * (r[1] - q[1])
 
#     if val == 0:
#         return 0
#     elif val > 0:
#         return 1
#     else:
#         return 2

# def convex_hull(points):
#     """
#     Input: 
#         points: an array of coordinates, i.e. [(x1, y1), (x2, y2), (x3, y3), ....]
#     Output: 
#         lst: a smaller array of coordinates, now only contain vertices of the hull with no duplicates
#     """
#     # YOUR CODE HERE
    
#     if len(points) < 3:
#         return
 
#     l = l_index(points)
 
#     hull = []
    
#     p = l
#     q = 0
    
#     while(True):
#         hull.append(p)
 
#         q = (p + 1) % len(points)
 
#         for i in range(len(points)):
#             if(orientation(points[p], points[i], points[q]) == 2):
#                 q = i
 
#         p = q
 
#         if(p == l):
#             hull.append(p)
#             break

#     final = []
    
#     for each in hull:
#         final.append(points[each])
    
#     zip(*points)
#     zip(*final)
#     plt.plot(*zip(*final))
#     plt.scatter(*zip(*points))
#     plt.show()
    
#     final.pop()
    
#     return final


def l_index(points):
    min = 0
    for i in range(len(points)):
        if points[i][0] <= points[min][0]:
            min = i
    return min

def orientation(p, q, r):
    val = (q[1] - p[1]) * (r[0] - q[0]) - (q[0] - p[0]) * (r[1] - q[1])
    
    if val == 0:
        return 0
    elif val > 0:
        return 1
    else:
        return 2

def convex_hull(points):
    
    if (len(points) < 3):
        return
    
    hull = []
    
    l = l_index(points)
    
    p = l
    q = 0
    
    while (True):
        hull.append(p)
        
        q = (p + 1) % len(points)
        
        for i in range(len(points)):
            if (orientation(points[p], points[i], points[q]) == 2):
                q = i
        
        p = q
        
        if (p == l):
            break
            
    final = []
    
    for each in hull:
        final.append(points[each])
        
    return final
                
    raise NotImplementedError()
```

<!-- grade_id: cell-1e128c225446bbb0, points: 4, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell

assert sorted(convex_hull([(4, 82), (80, 43), (74, 67), (77, 83), (60, 41), (4, 6), (38, 23), 
                    (4, 19), (92, 70), (62, 74), (84, 30), (63, 22), (37, 45), (56, 81), 
                    (31, 19), (31, 0), (91, 48), (34, 22), (58, 72), (22, 78), (78, 28), 
                    (15, 15), (49, 89), (13, 39), (11, 68), (18, 28), (49, 61), (11, 35), 
                    (15, 6), (61, 22), (63, 82), (31, 39), (1, 50), (50, 74), (77, 67), 
                    (58, 83), (46, 42), (62, 24), (66, 39), (93, 11), (0, 47), (29, 94), 
                    (2, 60), (43, 83), (67, 97), (32, 59), (97, 73), (43, 72), (33, 85), 
                    (48, 87), (41, 48), (95, 0), (98, 28), (17, 70), (55, 26), (7, 23), 
                    (60, 31), (79, 89), (47, 2), (43, 34), (94, 84), (12, 92), (47, 69), 
                    (25, 67), (95, 90), (16, 31), (5, 58), (84, 71), (15, 6), (23, 12), 
                    (57, 48), (36, 59), (84, 61), (95, 18), (16, 66), (84, 94), (2, 5), 
                    (9, 5), (97, 77), (0, 58), (37, 17), (25, 70), (51, 42), (78, 63), 
                    (42, 70), (31, 63), (93, 9), (33, 10), (2, 46), (66, 68), (74, 57), 
                    (69, 93), (86, 76), (78, 73), (55, 22), (68, 41), (69, 43), (97, 5), 
                    (95, 22), (27, 89)])) == sorted([(0, 47), (2, 5), (31, 0), (95, 0), (97, 5), 
                                             (98, 28), (97, 77), (95, 90), (84, 94), (67, 97), 
                                             (29, 94), (12, 92), (4, 82), (0, 58)])
print('All Tests Passed!!!')


```

Output:
```
All Tests Passed!!!
```

```python

```