```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import pandas as pd
import numpy as np
import numpy.testing as nt
```

# Optimisation II - Lab 2


## Instructions

* Carefully read all the instructions.
* Do **NOT** change the file name/format upon submission. Please do not add or remove cells from the notebook. 

* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!


----
This lab material has been created by Dr Matthew Woolway and has been edited by Krupa Prag. 

## Warm-up Exercises
Complete the following warm-up tasks:

----

### Question 1 (1 Mark)

Write a function that takes as its input argument a scalar integer named $\texttt{n}$. The function returns $\texttt{Q}$, a 2n-by-2n matrix. Q consists of four $n$-by-$n$ submatrices. The elements of the submatrix in the top left corner are all 1s, the elements of the submatrix at the top right are 2s, the elements in the bottom left are 3s, and the elements in the bottom right are 4s. 

<!-- grade_id: cell-af6aef1856fc920d, points: 0, grade: false, solution: true -->
```python
def matrix_quad(n):
    """
    Inputs:
        n: a scalar integer
    Outputs:
        mat: a numpy matrix
    """
    # YOUR CODE HERE
    #-----------Attempt 1-----------#
    # arr = np.ones((2*n, 2*n))
    # arr = np.append(arr2, arr3) 
    # for i in range(2*n):
    #     for j in range(2*n):
    #         if i < n:
    #             if j < n:
    #                 arr[i,j] = 1 
    #             if j >= n:
    #                 arr[i,j] = 2  
    #         if i >= n:
    #             if j < n:
    #                 arr[i,j] = 3 
    #             if j >= n:
    #                 arr[i,j] = 4 
    #------------------------------#
    arr1 = np.full((n, 2*n), [1] * n + [2] * n)
    arr2 = np.full((n, 2*n), [3] * n + [4] * n)
    arr = np.append(arr1, arr2, axis=0)
               
    return arr
    raise NotImplementedError()
```

<!-- grade_id: cell-e336471e068664ca, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
test1 = np.array([[1., 1., 2., 2.],
                 [1., 1., 2., 2.],
                 [3., 3., 4., 4.],
                 [3., 3., 4., 4.]])
nt.assert_array_almost_equal(test1, matrix_quad(2), err_msg="Function Incorrect")
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 2 (1 Mark)

Write a function that takes as inputs two positive integer scalars, $\texttt{n}$ and $\texttt{m}$, in that order. The function must create and return $board$, which is an $n$-by-$m$ matrix. Every element of $board$ is either 0 or 1. The first element, $board(0,0)$ is 1. No direct neighbours in the matrix, vertically or horizontally, can be equal. That is, a 1 element cannot have 1 immediately preceding or following it in the same row or column. 
(You can this of this as a checked board where 0 and 1 represent the black and white blocked of the board.)

<!-- grade_id: cell-013ce85b3036bcf5, points: 0, grade: false, solution: true -->
```python
def check_board(n, m):
    """
    Inputs:
        n, m: scalar integers
    Outputs:
        board: binary numpy matrix
    """
    # YOUR CODE HERE
    rE = np.zeros((1, m))
    rO =np.ones((1, m))
    
    for i in range(0,m,2):
        rE[0,i] = 1
        rO[0,i] = 0 

    if n % 2 == 0:
        arr = np.row_stack(n//2*(rE, rO))
    else:
        arr = np.row_stack((n//2)*(rE, rO))
        arr = np.append(arr, rE, 0)
            
    return arr
    raise NotImplementedError()
```

<!-- grade_id: cell-76467ce0faa7e738, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
test1 = np.array([[1., 0., 1., 0., 1.],
                  [0., 1., 0., 1., 0.],
                  [1., 0., 1., 0., 1.],
                  [0., 1., 0., 1., 0.]])
nt.assert_array_almost_equal(test1, check_board(4, 5), err_msg="Function Incorrect")
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 3 (1 Mark)

Write a function that takes three input arguments: $\texttt{limit}$, $\texttt{n}$, and $\texttt{m}$, in that order. The function returns an $n$-by-$m$ matrix of uniformly distributed random integers between 1 and limit inclusive.  You are NOT allowed to use the built-in function `randi`, but you can use `rand`. You will also find the built-in function `floor` useful for obtaining integers. 

As an enrichment exercise to this question, to ensure that your result is indeed uniformly distributed, test the output of the function by using the built-in function `hist`, which plots a histogram on larger inputs.

<!-- grade_id: cell-2318dc09deefd2ef, points: 0, grade: false, solution: true -->
```python
def random_dist(limit, n, m):
    # YOUR CODE HERE
    arr = np.random.rand(n,m)
    
    for i in range(n):
        for j in range(m):
            arr[i,j] *= (limit+1)
            
            if arr[i,j] < 1:
                arr[i,j] += 1
            
            arr[i,j] = int(arr[i,j])
            
    #plt.hist(arr)
    return arr 

    raise NotImplementedError()
```

<!-- grade_id: cell-4dbeddc62c1438c4, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
test1 = random_dist(5, 5, 5)
def isinteger(x):
     return np.equal(np.mod(x, 1), 0)
assert(np.all(isinteger(test1)) == True)
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

# Main Exercises

### Question 1 (5 Marks)

Write a function that implements the Secant method. It should take as inputs the function handle $\texttt{f'}$ (the first derivative of the function $\texttt{f}$), two initial guesses $x0, \ x1$ and a tolerance $\texttt{tol}$. The function should return the root of the function. Use the absolute difference to control your tolerance.

<!-- grade_id: cell-dfe21f59e51fcb1a, points: 0, grade: false, solution: true -->
```python
def secant_min(f, x0, x1, tol):
    """
    Inputs:
        f : a function handle
        x0: an initial point - scalar
        x1: a subsequent point - scalar
        
    Outputs:
        x: root of function f - scalar
    """
    # YOUR CODE HERE
    while (abs(x1 - x0) > tol):
        xn = x1 - f(x1) * (x1 - x0) / (f(x1) - f(x0))
        x0 = x1
        x1 = xn
    
    return xn
    raise NotImplementedError()
```

<!-- grade_id: cell-f648fc1436c9207f, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
f   = lambda x: x**2 - 612
x0  = 10
x1  = 30
tol = 10**(-5)
nt.assert_almost_equal(24.738633753705965, secant_min(f, x0, x1, tol), err_msg='incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 2 (7 Marks)

Write a function that performs the Golden Search method. It should take in the function handle $\texttt{f}$, the domain, $\texttt{a}$, $\texttt{b}$ and a iterations $\texttt{iter}$. It should return a both the minima (that is the midpoint between the interval) and the function value at the minima at the end of a set number of iterations. 

<!-- grade_id: cell-a5da4c9ae3fd7e69, points: 0, grade: false, solution: true -->
```python
def golden_search(f, a, b, iter):
    """
    Inputs:
        f : a function handle
        a : left endpoint - scalar
        b : right endpoint - scalar
        iter: set number of iterations - scalar
        
    Outputs:
        out: np.array of xmin and f(xmin)
    """
    # YOUR CODE HERE
    R = (sqrt(5)-1)/2
    
    for i in range(iter):
      x1 = a + R*(b-a)
      x2 = b - R*(b-a)
      print(x1)
      print(x2)
      
      if f(x1) > f(x2):
        b = x1
        x1 = x2
        x2 = b - R*(b-a)
        
      elif f(x1) < f(x2):
        a = x2
        x2 = x1
        x1 = a + R*(b-a)
      
    
    return [(b+a)/2, f((b+a)/2)]
    raise NotImplementedError()
```

<!-- grade_id: cell-9417231567d45acf, points: 7, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
f = lambda x: x**2 -6*x +15
a = 0
b = 10
iter = 4
nt.assert_array_almost_equal(np.array([ 0.72949017, 11.15521489]), golden_search(f, a, b, iter), err_msg='incorrect function')
print('All Tests Passed!!!')
```

Output:
```
0.6180339887498949
0.3819660112501051
0.3819660112501052
0.2360679774997897
0.47213595499957944
0.38196601125010515
0.5278640450004206
0.4721359549995794
    [... skipping hidden 1 frame]

File c:\Python310\lib\site-packages\numpy\testing\_private\utils.py:844, in assert_array_compare(comparison, x, y, err_msg, verbose, header, precision, equal_nan, equal_inf)
    840         err_msg += '\n' + '\n'.join(remarks)
    841         msg = build_err_msg([ox, oy], err_msg,
    842                             verbose=verbose, header=header,
    843                             names=('x', 'y'), precision=precision)
--> 844         raise AssertionError(msg)
    845 except ValueError:
    846     import traceback

AssertionError: 
Arrays are not almost equal to 6 decimals
incorrect function
Mismatched elements: 2 / 2 (100%)
Max absolute difference: 11.81420053
Max relative difference: 17.92785733
 x: array([ 0.72949 , 11.155215])
 y: array([ 0.454915, -0.658986])
```

```python

```