```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
from numpy import linalg as LA
from functools import reduce

```

# Optimisation II - Lab 3


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!
* You may not import any additional libraries, please use only those provided.
    
Developed by Dr M. Woolway and edited by Krupa Prag. 

## Warm-up Exercises



### Question 1 [4 Marks]

Write a function that takes in three scalar inputs and returns these values in a single row vector of increasing order. The function may not use any builtin functions but must perform this task only using if statements (possibly nested).


<!-- grade_id: cell-34cd77183ff3b4dd, points: 0, grade: false, solution: true -->
```python
def tri_sort(a, b, c):
    # YOUR CODE HERE
    ordered = []
    
    if (a < b and a < c):
        ordered.append(a)
        if (b < c):
            ordered.append(b)
            ordered.append(c)
        else:
            ordered.append(c)
            ordered.append(b)
    elif (b < a and b < c):
        ordered.append(b)
        if (a < c):
            ordered.append(a)
            ordered.append(c)
        else:
            ordered.append(c)
            ordered.append(a)
    else:
        ordered.append(c)
        if (b < a):
            ordered.append(b)
            ordered.append(a)
        else:
            ordered.append(a)
            ordered.append(b)
    
    return ordered
    raise NotImplementedError()
```

<!-- grade_id: cell-80fba3d18235871a, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark
nt.assert_array_equal(np.array([1, 5, 8]), tri_sort(5, 8, 1))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-cfb85bbc7f0b94c4, points: 3, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# Hidden cases
# 3 Marks

```

### Question 2 [4 Marks]

Write a function that takes in one input argument, $x$. If $x$ is an empty matrix, the function must return -1. If $x$ is a scalar it must return $0$. If $x$ is a vector it must return 1. Finally, if $x$ is none of these it must return 2. You may not use any builtin function such as `ðš’ðšœðšœðšŒðšŠðš•ðšŠðš›` etc.


<!-- grade_id: cell-de763a3b2eaf163b, points: 0, grade: false, solution: true -->
```python
def classifier(x):
    # YOUR CODE HERE
    try:
        if(x.size == 0):
            return -1
        else:
            try:
                x[1].size
                return 2
            except:
                return 1
    except:
        if (len(x) == 1):
            try:
                x[0] + 1
                return 0
            except:
                return 2
        return 2
    
    raise NotImplementedError()
```

<!-- grade_id: cell-867ca2a6eada40a7, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark3
assert(2 == classifier(np.random.rand(2, 5)))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-979237836c4d432a, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark
assert(1 == classifier(np.random.rand(1, 5)))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-b31a796bae990b57, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark
assert(-1 == classifier(np.array([])))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-ad828bad1788aeb0, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark
assert(0 == classifier([1]))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

## Main Exercises

Complete the following problems.


## Question 1 [6 Marks]
### Directional derivative of maximal rate of increase/decrease at a point

The function ``directional_derivative`` returns the rate of maximal increase or decrease of the function $\mathbf{f}$ at a given point $\mathbf{p}$. 

The function inputs are the function $\mathbf{f}$, the grad $\mathbf{g}$ ($\mathbf{\nabla f}$), the point $\mathbf{p}$, $n$ the dimension/number of variables that make up the function and a binary value which indicates maximal increase or decrease.

* $\mathbf{f}$ a function of $x_1, x_2, ..., x_n$.
* $\mathbf{g} = [g_1, g_2, ..., g_n]^T$, where $g_1, g_2, ..., g_n$ are respectively the grad of some function $f$ with respect to $x_1, x_2, ..., x_n$. This is a list. 
* $\mathbf{p} = (x_1, x_2, ..., x_n)$ which is a tuple.
* $n$ is the number of variables.
* $m$ is a binary value. 0: if maximal rate of decrease and 1: if maximal rate of increase required.

The function returns the maximal rate of increase/decrease at the the point $\mathbf{p}$. 


<!-- grade_id: cell-c5c9826f180ed489, points: 0, grade: false, solution: true -->
```python
def directional_derivative(f, g, p, n, m):
    """
    INPUTS:
        f: Function
        g: Grad of function f. Note that this is given in a vector format of dimensions nX1.
        p: Point given as an one-dimensional array
        n: number of variables
        m: binary value indicating maximal rate of decrease/increase
    OUTPUT: 
        maximal directional derivative of increase/decrease at the point p
    """
    # YOUR CODE HERE
    buffer = [func(*p)**2 for func in g]
    dd = sum(buffer)
      
    if m == 0:
        return -sqrt(dd)
    
    return sqrt(dd)
    raise NotImplementedError()
    
```

<!-- grade_id: cell-b73a99568f61a453, points: 6, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 7 Marks
f = lambda a, b: a**2 + 4*(b**2) + 6*a*b + 3*a + 4*b
g1 = lambda a, b: 2*a + 6*b +3
g2 = lambda a,b: 16*b + 6*a +4
g = np.transpose([g1, g2])
p = (1,1)
n = 2
m = 1
nt.assert_almost_equal(np.sqrt(797), directional_derivative(f, g, p, n, m), err_msg='incorrect function')
print('All Tests Passed!!!')

```

Output:
```
All Tests Passed!!!
```

### Question 2 [6 Marks]

Write a function `fixed_steepest_descent` which implements the method of Steepest Descent using a fixed step size. The function should take as inputs the multivariate function `f`, the gradient `g`, some initial guess `x`, the fixed step size `alpha` and some tolerance `tol`. Use `norm(x, 2)` to control your tolerance. The function should return the approximation to the minimum as well as the number of iterations required to find it.

<!-- grade_id: cell-7f4dbfed0c4f0a15, points: 0, grade: false, solution: true -->
```python
def fixed_steepest_descent(f, g, x, alpha, tol):
    # YOUR CODE HERE
    x_old = np.array([0, 0]) 
    x_list, y_list = [x], [f(x)]
    
    while np.linalg.norm(x,2) > tol:
        x_old = x
        s_k = g(x_old)
        x = x_old + alpha * s_k
        x_list.append(x)
        y_list.append(f(x))
    
    return [x, len(x_list)]
    raise NotImplementedError()
```

<!-- grade_id: cell-5df8a18d31c21c02, points: 6, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 6 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([-4*x[0], -6*x[1]])
x0    = np.array([1, 1]) 
alpha = 52/280
tol   = 1e-5

nt.assert_array_almost_equal([ 4.91552813e-06, -3.32604042e-09], fixed_steepest_descent(f1, g1, x0, alpha, tol)[0])
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

Please ensure your file name is ``'lab3.ipynb'``

```python
def dfp(f, g, x0, tol, max_iter, alpha):
    # YOUR CODE HERE
    k = 0
    G = np.eye(len(x0))
    
    while np.linalg.norm(g(x0),2) > tol and k < max_iter:
        d = np.matmul(-G, g(x0))
        x1 = x0 + alpha * d
        dx = x1 - x0
        dg = g(x1) - g(x0)
        
        delta = np.outer(dx,dx) / np.dot(dx,dg)
        
    k = 0
    G = np.eye(len(x0))

    while np.linalg.norm(g(x0), 2) > tol:
        d = np.matmul(-G, g(x0))
        x1 = x0 + alpha * d
        dx = x1 - x0
        dg = g(x1) - g(x0)

        delta = np.outer(dx, dx) / np.dot(dx, dg)
        gamma = np.outer(np.matmul(G, dg), np.matmul(dg, G)) / np.dot(np.matmul(dg, G),dg)

        G = G + (delta - gamma)

        x0 = x1
        k = k + 1

    return (x1, k)

    raise NotImplementedError()
```

```python
# Run this test cell to check your code
# Do not delete this cell
# 5 Marks
f = lambda x: x[0] - x[1] + 2*x[0]**2 + 2*x[0]*x[1] + x[1]**2
g = lambda x: np.array([1 + 4*x[0] + 2*x[1], -1 + 2*x[0] + 2*x[1]])
x0 = np.array([0, 0])
alpha = 0.75
tol = 1e-7
max_iter = 20
nt.assert_almost_equal([-1.0, 1.50000001], dfp(f, g, x0, tol, max_iter, alpha)[0])
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```