```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
from numpy import linalg as LA

```

# Optimisation II - Lab 4


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

## Question 1 [3 Marks]

Write a function `random_walk` which takes as inputs, a multivariate function $f$, an initial starting point $x$, a step length $l$, a tolerance for $l$, $tol$, the dimension of $f$, $n$, and a maximum number of iterations $it$. The function should return the approximation to the minima.


<!-- grade_id: cell-9bc29e7f371c486f, points: 0, grade: false, solution: true -->
```python
def random_walk(f, x, l, eps, n, it):
    """
    Inputs:
        f : a function handle
        x : initial search point
        l : step-length. Note that this is a list of alpha values of the same length as $i$. 
        n : dimension of $f$Aaa
        tol : tolerance threshold for step-length. 
        it : number of iterations to be performed.

    Outputs:
        x : updated x once method is terminated.
    """


    # YOUR CODE HERE
    f1 = f(*x)
    i = 1
    
    while l > eps:
        u = 2
        while u > 1:
            v = [np.random.uniform(-1,1) for j in range(n)]
            u = 1 / np.linalg.norm(v)
        
        uv = v / u
        print('this is v', v)
        xn = x + l * uv
        fn = f(*xn)
        
        if fn < f1:
            x = xn
            f1 = fn
            i = 1
        else:
            if i <= it:
                i += 1
            else:
                print('this is off')
                l = l / 2
    
    print(x)
    return x     
    
    
    # f1 = f(*x)          #2
    # i = 1               #3
    
    # while l > eps:
    #     u = 2           #4 ensuring no bias
    #     while u > 1: 
    #         v = [np.random.uniform(-1,1) for i in range(n)]         
    #         u = 1/np.linalg.norm(v) 
            
    #     uv = v / np.linalg.norm(v) #5
    #     xn = x + l*uv
    #     fn = f(*xn)
        
    #     if fn < f1:     #6
    #         x = xn
    #         f1 = fn
    #         i = 1
    #     else:           #7
    #         if i <= it:
    #             i += 1    
    #         elif i > it:
    #             l = l/2 #8
            
    # return x   
    raise NotImplementedError()
```

<!-- grade_id: cell-79ddddae7f98d53e, points: 3, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 3 Marks
f   = lambda x1, x2: x1 - x2 + 2*x1**2 + 2*x1*x2 + x2**2
x0  = np.array([0, 0])
lam = 1
eps = 1e-5
n   = 2 
N   = 200
nt.assert_almost_equal(-1.00000726, random_walk(f, x0, lam, eps, n, N)[0], decimal=4)
print('Test case passed!!!')
```

Output:
```
this is v [-0.8414811691640267, 0.713410472325978]
this is v [0.4433621343567171, 0.9651463475310038]
this is v [0.9550917191036405, -0.7511580914929235]
this is v [0.5118165683050515, -0.8733092783441974]
this is v [0.8085018888746884, -0.6531965750003572]
this is v [0.978173563852009, -0.3812158620267443]
this is v [-0.8370256894596326, 0.5726923003910362]
this is v [-0.9347927086893342, 0.4392665250058929]
this is v [-0.7401815244074468, 0.9827427773775428]
this is v [-0.8060620168472277, 0.5953864039988184]
this is v [0.9200745072233825, -0.9376162121650407]
this is v [-0.8459155531730282, -0.7744599865114097]
this is v [0.6813773065701854, 0.9792170784947996]
this is v [-0.920746517299152, -0.7087897664021847]
this is v [0.18577541913179507, -0.9960053721590036]
this is v [0.8391706904342895, -0.8038226644944473]
this is v [-0.6010332684358175, 0.9534269886849578]
this is v [0.3268184159207097, -0.9780039685223727]
this is v [-0.9903710390293459, 0.8081714938499398]
this is v [-0.
```

## Question 2 [7 Marks]

Write a function to perform the Downhill Simplex Method. The function `downhill_simplex` takes as inputs a multivariate function $f$, an initial guess $x0$, a maximum number of iterations $it$ along with values for $alpha$ - relection parameter, $gamma$ - expansion parameter and $beta$ - contraction parameter. The function should return the minimum of the function $f$.



<!-- grade_id: cell-8082c577150ea5e9, points: 0, grade: false, solution: true -->
```python
def downhill_simplex(f, x0, it, alpha, gamma, beta):
    """
    Inputs:
        f : a function handle
        x0: initial point  
        it : number of iterations to be performed.
        alpha: reflection paramater.
        gamma: expansion paramater.
        beta: contraction paramater.


    Outputs:
        x : updated x once method is terminated.
    """    
    # YOUR CODE HERE        
    n = len(x0)
    
    for i in range(it):
        x = []
        
        for j in sorted([f(j) for j in x0]):
            [x.append(list(k)) for k in x0 if j == f(k)]
            
        x = np.array(x)
        G = np.mean(x[:n-1:], 0)
        
        xr = (1 + alpha) * G - alpha * x[n-1]
        
        if f(x[0]) <= f(xr) and f(xr) <= f(x[n-2]):
            x[n-1] = xr
            
        elif f(xr) < f(x[0]):
            xe = gamma * xr + (1 - gamma) * G
            if (f(xe) < f(x[0])):
                x[n-1] = xe
            else:
                x[n-1] = xr
                
        elif f(xr) > f(x[n-2]):
            if f(xr) < f(x[n-1]):
                xc = beta * xr + (1 - beta) * G
            else:
                xc = beta * x[n-1] + (1 - beta) * G
                
            if f(xc) < f(x[n-1]) and f(xc) < f(xr):
                x[n-1] = xc
            else:
                x = x/2
                
        x0 = x
        
    return x
    

    raise NotImplementedError()

```

<!-- grade_id: cell-071545b6489e7391, points: 7, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 7 Marks
x0    = np.array([[-1.20, 1.00], [-0.23, 1.26], [ -0.94, 1.97]])
f     = lambda x: (1 - x[0])**2 + 1000*(x[1] - x[0]**2)**2
alpha = 1
beta  = 0.5
gamma = 2
it    = 200
t1 = np.array([[1., 1.],
               [1., 1.],
               [1., 1.]])
nt.assert_almost_equal(t1, downhill_simplex(f, x0, it, alpha, gamma, beta))
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

```python

```